# Vamos começar

Tópicos:
1. Como imprimir
2. Como atribuir variáveis
3. Como comentar
4. Sintaxe para matemática básica

## Como imprimir

Na linguagem Julia, geralmente usamos `println ()` para imprimir

```julia
println("Estou animado para aprender Julia!")
```

Output:
```
Estou animado para aprender Julia!
```

```julia
123
456; # ponto e virgula sem output 
```

## Como atribuir variáveis

Tudo o que precisamos é de um nome para a variável, um valor e um sinal de igual!<br>
Julia vai computar os tipos (types) para nós.

```julia
my_answer = 42
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

```julia
😺 = "gato sorrindo!"
typeof(😺)
```

Output:
```
String
```

Para o emoji gato sorrindo, use o preenchimento da guia para selecionar o nome do emoji e, em seguida, pressione novamente

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
\
```

Output:
```
syntax: "\" is not a unary operator
```

Depois de atribuir um valor a uma variável, podemos reatribuir outro valor de um tipo diferente a essa variável sem qualquer problema.

```julia
😺=1
```

Output:
```
1
```

```julia
typeof(😺)
```

Output:
```
Int64
```

Observação: Julia nos permite escrever um código genérico, o emoji 😺 é um exemplo disso. 

Isso nos permite escrever um código como por exemplo,

```julia
😀 = 0
😞 = -1
```

Output:
```
-1
```

```julia
😺 + 😞 == 😀
```

Output:
```
true
```

## Como comentar

```julia
# Você pode deixar comentários em uma única linha usando a tecla /hash (#)
```

```julia
#=

Para várias linhas de comentários, 
usar '#= =#'.

=#
```

## Sintaxe para matemática básica

```julia
sum = 3 + 7
```

Output:
```
10
```

```julia
difference = 10 - 3
```

Output:
```
7
```

```julia
product = 20 * 5
```

Output:
```
100
```

```julia
quotient = 100 / 10
```

Output:
```
10.0
```

```julia
power = 10 ^ 2
```

Output:
```
100
```

```julia
modulus = 101 % 2
```

Output:
```
1
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia
?convert
```

Output:
```
50
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia
days = 365
convert(Float64, days)
days_float = days
```

Output:
```
365
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
?parse
```

Output:
```
search: [0m[1mp[22m[0m[1ma[22m[0m[1mr[22m[0m[1ms[22m[0m[1me[22m try[0m[1mp[22m[0m[1ma[22m[0m[1mr[22m[0m[1ms[22m[0m[1me[22m [0m[1mp[22m[0m[1ma[22m[0m[1mr[22mtial[0m[1ms[22mortp[0m[1me[22mrm [0m[1mp[22m[0m[1ma[22m[0m[1mr[22mtial[0m[1ms[22mortp[0m[1me[22mrm! [0m[1mp[22m[0m[1ma[22mi[0m[1mr[22m[0m[1ms[22m ski[0m[1mp[22mch[0m[1ma[22m[0m[1mr[22m[0m[1ms[22m

[36m  parse(type, str; base)[39m

  Parse a string as a number. For [36mInteger[39m types, a base can be specified (the
  default is 10). For floating-point types, the string is parsed as a decimal
  floating-point number. [36mComplex[39m types are parsed from decimal strings of the
  form [36m"R±Iim"[39m as a [36mComplex(R,I)[39m of the requested type; [36m"i"[39m or [36m"j"[39m can also be
  used instead of [36m"im"[39m, and [36m"R"[39m or [36m"Iim"[39m are also permitted. If the string
  does not contain a valid number, an error is ra
```

```julia
Please click on `Validate` on the top, once you are done with the exercises.
```

Output:
```
syntax: extra token "click" after end of expression
```

```julia

```