# Strings

Tópicos:
1.  Como obter uma string
2. Interpolação de string
3. Concatenação de string

## Como obter uma string

Coloque os seus caracteres em " " ou """ """!

```julia
s1 = "Eu sou uma string."
```

Output:
```
"Eu sou uma string."
```

```julia
s2 = """Eu também sou uma string. """
```

Output:
```
"Eu também sou uma string. "
```

"Aqui, obtemos um "erro" porque é ambíguo onde essa string termina"

```julia
"""Olha mamãe, sem "erros"!!! """
```

Output:
```
"Olha mamãe, sem \"erros\"!!! "
```

Note que ' ' define um carácter, mas NÃO uma string!

```julia
typeof('a')
```

Output:
```
Char
```

```julia
'Vamos obter um erro aqui'
```

Output:
```
@ .\boot.jl:360 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base .\loading.jl:1094
```

## Interpolação de strings

Podemos utilizar o sinal $ para inserir variáveis existentes numa string e avaliar expressões dentro de uma string. <br>
Abaixo está um exemplo que contém alguma informação pessoal.

```julia
nome = "Joana"
num_dedosmao = 10
num_dedospe = 10
```

Output:
```
10
```

```julia
println("Olá, meu nome é $nome.")
println("Eu tenho $num_dedosmao dedos da mão e $num_dedospe dedos do pé.")
```

Output:
```
Olá, meu nome é Joana.
Eu tenho 10 dedos da mão e 10 dedos do pé.
```

```julia
 println("Ou seja, $(num_dedosmao + num_dedospe) dedos no total!!")
```

Output:
```
Ou seja, 20 dedos no total!!
```

## Concatenação de string

Abaixo estão três formas de concatenar strings! <br><br>
A primeira maneira é utilizar a função `string()`. <br>
`string()` converte entradas sem string em strings.

```julia
s3 = "Quantos gatos ";
s4 = "são muitos gatos?";
😺 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"Quantos gatos são muitos gatos?"
```

```julia
string("Não sei, mas ", 😺, " é muito pouco.")
```

Output:
```
"Não sei, mas 10 é muito pouco."
```

Também podemos utilizar `*` para concatenação!

```julia
s3*s4
```

Output:
```
"Quantos gatos são muitos gatos?"
```

### Exercícios

#### 2.1 
Criar uma string que diga "oi" 1000 vezes, primeiro com "repetição" e depois com o operador de exponenciação, que pode chamar `*` abaixo. Atribuir-lhe a variável `oi` abaixo.

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia
oi = repeat("oi", 1000);
```

```julia
oi2 = "oi"^1000;
```

```julia
@assert oi == "oioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioioi"
```

#### 2.2 
Declare duas variáveis

```julia
a = 3
b = 4
```
use-as para criar duas strings:
```julia
"3 + 4"
"7" 
```
e armazene os resultados em `c` and `d` respectivamente

```julia
?@assert
```

Output:
```
[36m  @assert cond [text][39m

  Throw an [36mAssertionError[39m if [36mcond[39m is [36mfalse[39m. Preferred syntax for writing
  assertions. Message [36mtext[39m is optionally displayed upon assertion failure.

[33m[1m  │ [22m[39m[33m[1mWarning[22m[39m
[33m[1m  │[22m[39m
[33m[1m  │[22m[39m  An assert might be disabled at various optimization levels. Assert
[33m[1m  │[22m[39m  should therefore only be used as a debugging tool and not used for
[33m[1m  │[22m[39m  authentication verification (e.g., verifying passwords), nor
[33m[1m  │[22m[39m  should side effects needed for the function to work correctly be
[33m[1m  │[22m[39m  used inside of asserts.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> @assert iseven(3) "3 is an odd number!"[39m
[36m  ERROR: AssertionError: 3 is an odd number![39m
[36m  [39m
[36m  julia> @assert isodd(3) "What even are numbers?"[39m
```

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia
a = 3
b = 4
c = string(a," + ", b)
d = string(a+b);
```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3 + 4"
@assert d == "7"
```

```julia
Por favor clique em `Validate` no topo, uma vez terminados os exercícios.
```