# Estruturas de dados

Quando começarmos a trabalhar com muitos dados ao mesmo tempo, será conveniente para nós armazenar dados em estruturas como arrays ou dicionários (em vez de confiarmos apenas em variáveis).<br>

Tipos de estruturas de dados abrangidos:
1. Tuples
2. Dicionários
3. Arrays

<br>
Como visão geral, os tuples e os arrays são sequências ordenadas de elementos (para que possamos indexar neles). Os dicionários e os arrays são ambos mutáveis.
Explicaremos isto mais abaixo!

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)
```


```julia
meusanimaisfav = ("pinguins", "gatos", "Petauro")
```

Output:
```
("pinguins", "gatos", "Petauro")
```

Podemos indexar neste tuple,

```julia
meusanimaisfav[1]
```

Output:
```
"pinguins"
```

mas como os tuples são imutáveis, não podemos atualizá-los

```julia
meusanimaisfav[1] = "otters"
```

Output:
```
@ .\boot.jl:360 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base .\loading.jl:1094
```

## Agora em 1.0: NamedTuples

Como podemos adivinhar, `NamedTuple`s são exatamente como `Tuple`s exceto que cada elemento adicionalmente tem um nome! Eles têm uma sintaxe especial usando `=` dentro de um tuple:

```julia
(nome1 = item1, nome2 = item2, ...)
```

```julia
meusanimaisfav = (passaro = "pinguins", mamifero = "gato", marsupial = "Petauro")
# isso parece um dicionário ordenado mas imutável, ou não ?
```

Output:
```
(passaro = "pinguins", mamifero = "gato", marsupial = "Petauro")
```

Como os regulares `Tuples`, `NamedTuples` são ordenados, para que possamos recuperar seus elementos através de indexação:

```julia
meusanimaisfav[1]
```

Output:
```
"pinguins"
```

Eles também acrescentam a habilidade especial de acessar valores por seu nome:

```julia
meusanimaisfav.passaro
```

Output:
```
"pinguins"
```

## Dicionários

Se tivermos conjuntos de dados relacionados uns aos outros, podemos optar por armazenar esses dados em um dicionário. Podemos criar um dicionário usando o `Dict()` que podemos inicializar como um dicionário vazio ou uma chave de armazenamento, pares de valores.

Syntax:
```julia
Dict(chave1 => valor1, chave2 => vslor2, ...)
```

Um bom exemplo é uma lista de contatos, onde associamos nomes a números de telefone.

```julia
listadetel = Dict("Joana" => "867-5309", "Caça-Fantasmas" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Caça-Fantasmas" => "555-2368"
  "Joana"          => "867-5309"
```

Neste exemplo, cada nome e número é um par "chave" e "valor". Podemos pegar o número da Joana (um valor) usando a chave associada

```julia
listadetel["Joana"]
```

Output:
```
"867-5309"
```

Podemos acrescentar outra entrada a este dicionário da seguinte forma

```julia
listadetel["Filipe"] = "555-FILK"
```

Output:
```
"555-FILK"
```

Vamos verificar como está a nossa lista telefônica agora...

```julia
listadetel
```

Output:
```
Dict{String, String} with 3 entries:
  "Caça-Fantasmas" => "555-2368"
  "Joana"          => "867-5309"
  "Filipe"         => "555-FILK"
```

Podemos excluir Filipe de nossa lista de contatos e simultaneamente obter seu número - usando `pop!`

```julia
pop!(listadetel, "Filipe")
```

Output:
```
"555-FILK"
```

```julia
listadetel
```

Output:
```
Dict{String, String} with 2 entries:
  "Caça-Fantasmas" => "555-2368"
  "Joana"          => "867-5309"
```

Ao contrário dos tuples e arrays, os dicionários não são encomendados. Portanto, não podemos indexar neles.

```julia
listadetel[1]
```

Output:
```
@ .\boot.jl:360 [inlined]
 [4] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base .\loading.jl:1094
```

No exemplo acima, `julia` acha que você está tentando acessar um valor associado à chave `1`.

## Arrays

Ao contrário dos tuples, os arrays são mutáveis. Ao contrário dos dicionários, os arrays contêm coleções ordenadas. <br>
Podemos criar um array, encerrando essa coleção em `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]
```


Por exemplo, podemos criar uma matriz para acompanhar meus amigos

```julia
amigos = ["Filipe", "Monalisa", "Ana", "Fernando", "Rodrigo"]
```

Output:
```
5-element Vector{String}:
 "Filipe"
 "Monalisa"
 "Ana"
 "Fernando"
 "Rodrigo"
```

O `1` no `Array{String,1}` significa que este é um vetor unidimensional. Um `Array{String,2}` seria uma matriz 2d, etc. A `String` é o tipo de cada elemento.

ou para armazenar uma sequência de números

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mixture = [1, 1, 2, 3, "Rodrigo", "Rubens"]
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Rodrigo"
  "Rubens"
```

Uma vez que tenhamos um array, podemos pegar pedaços individuais de dados de dentro desse array, indexando-os ao array. Por exemplo, se quisermos o terceiro amigo listado em `amigos`, escrevemos

```julia
amigos[3]
```

Output:
```
"Ana"
```

Podemos usar a indexação para editar um elemento existente de um array

```julia
amigos[3] = "Francisco"
```

Output:
```
"Francisco"
```

Sim, Julia é uma indexação baseada em 1, não em 0 como Python.  As guerras são travadas por questões menores. Eu tenho um amigo com a sabedoria de Salomão que propõe resolver isto de uma vez por todas com ½ 😃

Também podemos editar o array usando as funções `push!` e `pop!. `push!` adiciona um elemento ao final de um array e `pop!` remove o último elemento do array

Podemos adicionar outro número à nossa seqüência fibonnaci

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
 21
```

e então remova-o

```julia
pop!(fibonacci)
```

Output:
```
21
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

Até agora eu dei exemplos de apenas arrays 1D de escalares, mas os arrays podem ter um número arbitrário de dimensões e também podem armazenar outros arrays. 
<br><br>
Por exemplo, os seguintes são arrays de arrays:

```julia
favoritos = [["kebab", "chocolate", "ovos"],["pinguins", "gatos", "pentauro"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["kebab", "chocolate", "ovos"]
 ["pinguins", "gatos", "pentauro"]
```

```julia
numeros = [[1, 2, 3], [4, 5], [6, 7, 8, 9]]
```

Output:
```
3-element Vector{Vector{Int64}}:
 [1, 2, 3]
 [4, 5]
 [6, 7, 8, 9]
```

Abaixo estão exemplos de arrays 2D e 3D povoada com valores aleatórios.

```julia
rand(4, 3)
```

Output:
```
4×3 Matrix{Float64}:
 0.881233  0.0557425  0.330672
 0.383833  0.344118   0.0662933
 0.632385  0.89101    0.344433
 0.563187  0.895307   0.349672
```

```julia
rand(4, 3, 2)
```

Output:
```
4×3×2 Array{Float64, 3}:
[:, :, 1] =
 0.0396961  0.967998  0.984955
 0.955625   0.294377  0.51831
 0.229512   0.764688  0.10264
 0.799889   0.676456  0.640694

[:, :, 2] =
 0.754194   0.763964  0.654917
 0.966966   0.736196  0.90908
 0.0712113  0.171226  0.523706
 0.682173   0.386637  0.313933
```

Seja cuidadoso quando quiser copiar arrays!

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mesmosnum = fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mesmosnum[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
 404
   1
   2
   3
   5
   8
  13
```

Editando `mesmosnum` faz com que `fibonacci` seja atualizado também!

No exemplo acima, na verdade não fizemos uma cópia de `fibonacci`. Acabamos de criar uma nova maneira de acessar as entradas no array vinculado a `fibonacci`.

Se quisermos fazer uma cópia do array vinculado a `fibonacci`, podemos usar a função `copy`.

```julia
# Primeiro, restaurar os fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mesmosnum = copy(fibonacci)
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
mesmosnum[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

Neste último exemplo, fibonacci não foi atualizado. Portanto, vemos que os arrays ligados a `mesmosnum` e `fibonacci` 
são distintos.

### Exercícios

#### 3.1 
Crie um array, `a_ray`, com o seguinte código:

```julia
a_ray = [1, 2, 3]
```

Adicione o número `4` ao final desta matriz e, em seguida, remova-o.

```julia
a_ray = [1, 2, 3];
push!(a_ray, 4);
println(a_ray)
pop!(a_ray);
println(a_ray)
```

Output:
```
[1, 2, 3, 4]
[1, 2, 3]
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Tente adicionar "Emergência" como chave para `listadetel` com o valor `string(911)` com o seguinte código
```julia
listadetel["Emergencia"] = 911
```

Por que isso não funciona?

```julia
listadetel["Emergencia"] = 911
```

Output:
```
@ .\boot.jl:360 [inlined]
 [4] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base .\loading.jl:1094
```

#### 3.3 

Crie um novo dicionário chamado `tel_flex` que tem o número da Joana armazenado como um número inteiro e o número dos Caça-Fantasmas armazenado como uma cadeia com o seguinte código

```julia
tel_flex = Dict("Joana" => 8675309, "Caça-Fantasmas" => "555-2368")
```

```julia
tel_flex = Dict("Joana" => 8675309, "Caça-Fantasmas" => "555-2368")
```

Output:
```
Dict{String, Any} with 2 entries:
  "Caça-Fantasmas" => "555-2368"
  "Joana"          => 8675309
```

```julia
@assert tel_flex == Dict("Joana" => 8675309, "Caça-Fantasmas" => "555-2368")
```

#### 3.4 
Acrescente a chave "Emergencia" com o valor `190` (um número inteiro) para `tel_flex`.

```julia
tel_flex["Emergencia"] = 190
```

Output:
```
190
```

```julia
@assert haskey(tel_flex, "Emergencia")
```

```julia
@assert tel_flex["Emergencia"] == 190
```

#### 3.5 
Por que podemos adicionar um número inteiro como um valor a `tel_flex` mas não a `listadetel`? Como poderíamos ter inicializado `listadetel` para que aceitasse inteiros como valores?

```julia
Dict{String, Int64}()
# Cria um dicionário vazio atribuindo a tipagem dos dados neste inicialização
```

Output:
```
Dict{String, Int64}()
```

Por favor, clique no botão `Validate` no topo, após a conclusão do exercício