# Funções

Topicos:
1. Como declarar uma função
2. Digitação de pato em Julia
3. Funções mutantes vs não mutantes
4. Algumas funções de ordem superior 

## How to declare a function
Julia nos dá algumas maneiras diferentes de escrever uma função. O primeiro requer as palavras-chave `function` e `end `

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
sayhi("Reginaldo")
```

Output:
```
Hi Reginaldo, it's great to see you!
```

```julia
function f(x)
    x^2
end
f(3)
```

Output:
```
9
```

Podemos chamar qualquer uma das funções como:

```julia
sayhi("C-3PO")
```

Output:
```
Hi C-3PO, it's great to see you!
```

```julia
f(42)
```

Output:
```
1764
```

Alternativamente, Poderíamos ter declarado qualquer uma dessas funções em uma única linha 

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

```julia
f2(x) = x^2
```

```julia
sayhi2("R2D2")
```

```julia
f2(42)
```

Finalmente, poderiamos ter declarado essas funções como "anônimas" 

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

```julia
f3 = x -> x^2
```

```julia
sayhi3("Chewbacca")
```

```julia
f3(42)
```

## Digitação de pato em Julia
*"Sem ele treme como um pato, é um pato."* <br><br>
As funções da Julia funcionaram em qualquer entrada que faça sentido. <br><br>
Por exemplo, `sayhi` trabalha no nome deste personagem da TV, escrito como um inteiro..

```julia
sayhi(55595472)
```

Output:
```
Hi 55595472, it's great to see you!
```

E trabalhará em uma matriz `f`. 

```julia
A = rand(3, 3)
A
```

Output:
```
3×3 Array{Float64,2}:
 0.929378   0.335597  0.232618
 0.0452895  0.321972  0.251592
 0.0524641  0.280963  0.351763
```

```julia
f(A)
```

Output:
```
3×3 Array{Float64,2}:
 0.891146   0.485306  0.38245
 0.0698725  0.189553  0.180041
 0.0799386  0.206901  0.206629
```

`f` também trabalhará em um string como "hi" porque é definido para entradas de string como concatenação de corda `*`.

```julia
f("hi")
```

Output:
```
"hihi"
```

Por outro lado, `f` não funcionará em um vetor. Ao contrário `A^2`, que é bem definido, o significado de um vetor `v^2` , `v`, não é uma operação algebrica bem definiida. 

```julia
v = rand(3)
```

Output:
```
3-element Array{Float64,1}:
 0.7219920508151683
 0.1237437754618207
 0.8670453876783517
```

```julia
f(v)
```

Output:
```
[2] literal_pow at ./none:0 [inlined]
 [3] f(::Array{Float64,1}) at ./In[4]:2
 [4] top-level scope at In[12]:1
```

## Funções mutantes vs não mutantes

Por convecção, as funções `!` alteram o seu conteúdo.

Por exmeplo, vamos olhar a diferença de  `sort` e `sort!`.


```julia
v = [3, 5, 2]
```

Output:
```
3-element Array{Int64,1}:
 3
 5
 2
```

```julia
sort(v)
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Array{Int64,1}:
 3
 5
 2
```

`sort(v)` retorna uma matriz classificada que contém os mesmos elementos que `v`, mas `v` é deixada iterada. <br><br>

Por outro lado, quando executamos `sort!(v)`, o conteúdo de  v é classidicado dentro da matriz `v`.

```julia
sort!(v)
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

## Algumas funções de ordem superior 

### mapa

`map` e uma função de  "higher-order" em Julia que *toma uma função* como um de seus argumento de entrada. 
em seguida, aplica essa função a cada elemento da estrutura de dados que você passa. Por examplo, executar `map` 

```julia
map(f, [1, 2, 3])
```
Lhe dará uma matriz de saída onde a função `f` foi aplicada a todos os elementos `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
 1
 4
 9
```

Aqui nós acertamos todos os elementos do vetor `[1, 2, 3]`, em vez do quadrado do vetor `[1, 2, 3]`.

Para fazer isso, podíamos ter passado para a função anônima em vez de uma função nomeada, como `map`

```julia
x -> x^3
```

Output:
```
#3 (generic function with 1 method)
```

via

```julia
map(x -> x^3, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
  1
  8
 27
```

E agora novamente todos os elementos de  `[1, 2, 3]`!

### Difusão

`broadcast` é uma função de ordem superior como `map`. `broadcast` é uma generalização de `map`, para que ele possa fazer tudo o que pode fazer e muito mais . A sintaxe para chamar `broadcast` é a mesma para chmar o `map`

```julia
broadcast(f, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
 1
 4
 9
```

E novamente, aplicamos `f` (ao quadrado) a todos os elementos de `[1, 2, 3]` - desta vez por "broadcasting" `f`!

Algum açúcar sintático para chamar `broadcast` é colocar `.` entre o nome da função que deseja e `broadcast` em seus argumentos de entrada. Por exemplo,

```julia
broadcast(f, [1, 2, 3])
```
e o mesmo que :
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
 1
 4
 9
```

Observe novamente como isso é diferente de chamar 
```julia
f([1, 2, 3])
```
Podemos acertar cada elemento de um vetor, mas não podemos acertar um vetor!

Para levar para casa o ponto, vamos olhar a diferença entre

```julia
f(A)
```
e
```julia
f.(A)
```
Para uma matriz `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

Output:
```
3×3 Array{Int64,2}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
f(A)
```

Output:
```
3×3 Array{Int64,2}:
  30   36   42
  66   81   96
 102  126  150
```

As before we see that for a matrix, `A`,
```
f(A) = A^2 = A * A
``` 

Por outro lado,

```julia
B = f.(A)
```

Output:
```
3×3 Array{Int64,2}:
  1   4   9
 16  25  36
 49  64  81
```

contem os quadrados de todas as entradas de `A`.

Esta sintaxe de ponto para trasmissão nos permite escrever expressões compostas relativamente complexas de forma que pareça natural, mais prócima da notação matemática. Por exemplo, podemos escrever 

```julia
A .+ 2 .* f.(A) ./ A
```

Output:
```
3×3 Array{Float64,2}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

em vez de

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

Output:
```
3×3 Array{Float64,2}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

E os dois vão exwcutar exatamente o mesmo.

### Exercícios

#### 6.1 
Escreva uma função que adicione 1 `a sua entrada.

```julia
add_one(xis) = xis+1
```

Output:
```
add_one (generic function with 1 method)
```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 6.2 
Use `map` ou `broadcast` para inclemetar cada elemento da matriz `A` e `1` é atribua a uma variavél `A1`.

```julia
A1 = broadcast(add_one, A)
```

Output:
```
3×3 Array{Int64,2}:
 2  3   4
 5  6   7
 8  9  10
```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 6.3 
Use a sintaxe `A1` de ponto `1` by para incrementar cada elemento da matriz e armazená-la em variável `A2`

```julia
A2=add_one.(A1)
```

Output:
```
3×3 Array{Int64,2}:
 3   4   5
 6   7   8
 9  10  11
```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

Por favor, clique em `Validate` em cima, uma vez que você tenha feito os exercícios.